# Objective

The objective of this notebook is to explore how to use notebooks for homework.
You will be submitting your homework during the course using these notebooks.

```python
STUDENT_NAME = "Nombre Apellido"
COURSE_NAME = "test1"
EXERCISE_NAME = "python-intro"
```

```python
import numpy as np
import pygradus as gr
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
/Users/danski/workspace/pygrade/examples/solutions/intro_python.ipynb Cell 3' in <cell line: 2>()
      <a href='vscode-notebook-cell:/Users/danski/workspace/pygrade/examples/solutions/intro_python.ipynb#ch0000004?line=0'>1</a> import numpy as np
----> <a href='vscode-notebook-cell:/Users/danski/workspace/pygrade/examples/solutions/intro_python.ipynb#ch0000004?line=1'>2</a> import pygradus as gr

ModuleNotFoundError: No module named 'pygradus'
```

<!-- grade_id: grade_vowels, points: 0, grade: false, solution: true -->
```python
def remove_vowels(s: str) -> str:
    """
    This function removes all vowels from strings
    Example: "hello" -> "hll"
    """
    
    #sss
    
    return s.translate(str.maketrans('', '', 'aeiouáéíóú'))
    
    #eee
```

<!-- grade_id: grade_remove_vowels, points: 1, grade: true, solution: false -->
```python
assert remove_vowels("hola") == "hl"
answer_1  = remove_vowels("este fue un buen día")
print(answer_1)
```

Output:
```
st f n bn d
```

<!-- grade_id: sum_of_list, points: 0, grade: false, solution: true -->
```python
def sum_list_of_floats(a: list, b:list) -> list:
    """
    Returns the coordinate wise sum of a and b.
    """
    
    #sss

    return list(np.array(a) + np.array(b))

    #eee
```

<!-- grade_id: grade_sum, points: 1, grade: true, solution: false -->
```python
assert np.allclose(
    sum_list_of_floats(
        [1, 2, 3], [4, 5, -2.3]),
    [5, 7, 0.7]
)
answer_2 = sum_list_of_floats([0, 0], [1, -1])
print(answer_2)
```

Output:
```
[1, -1]
```

```python
#ttt
TASKS = [
         {
             "name": "remove_vowels",
             "answer": str(answer_1),
         },
                  {
             "name": "sum_floats",
             "answer": str(answer_2),
         }
]
```

```python
#ddd
config = {
    "course_name": COURSE_NAME,
    "exercise_name": EXERCISE_NAME,
    "tasks": TASKS,
    
}
response = gr.create_exercise("admin", config)
```

Output:
```
··········
```

```python

```